Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Aseem Sachdeva"
```

---

# Information Visualization I 
## School of Information, University of Michigan

## Week 1: 
- Domain identification vs Abstract Task extraction
- Pandas Review 

## Assignment Overview
### The objectives for this week are for you to:

- Review, reflect, and apply the concepts of Domain Tasks and Abstract Tasks. Specifically, given a real context, identify the expert's goals and then abstract the visualization tasks. 

<img src="assets/domain-abstraction.png" alt="Drawing" style="width: 500px;"/>

- Review and evaluate the domain of [Pandas](https://pandas.pydata.org/) as a tool for reading, manipulating, and analyzing datasets in Python.

### The total score of this assignment will be 100 points consisting of:
- Case study reflection: Car congestion and crash rates (20 points)
- Pandas programming exercise (80 points)


### Resources:
- We're going to be recreating parts of this article by [CMAP](https://www.cmap.illinois.gov/) available [online](https://www.cmap.illinois.gov/updates/all/-/asset_publisher/UIMfSLnFfMB6/content/crash-scans-show-relationship-between-congestion-and-crash-rates) (CMAP, 2016)  
- We'll need the datasets from the city of Chicago. We have downloaded a subset to the local folder [/assets](assets/)
    - If you're curious, the original dataset can be found on [Chicago Data Portal](https://data.cityofchicago.org/)
        - [Chicago Traffic Tracker - Historical Congestion Estimates by Segment - 2011-2018](https://data.cityofchicago.org/Transportation/Chicago-Traffic-Tracker-Historical-Congestion-Esti/77hq-huss)
        - [Traffic Crashes - Crashes](https://data.cityofchicago.org/Transportation/Traffic-Crashes-Crashes/85ca-t3if)
- Altair
    - We will use a python library called [Altair](https://altair-viz.github.io/) for the visualizations. Don't worry about understanding this code. You will only need to prepare the data for the visualization in Pandas. If you do it correctly, our code will produce the visualization for you.

## Part 1. Domain identification vs Abstract Task extraction (20 points)
Read the following article by CMAP [Crash scans show the relationship between congestion and crash rates](https://www.cmap.illinois.gov/updates/all/-/asset_publisher/UIMfSLnFfMB6/content/crash-scans-show-relationship-between-congestion-and-crash-rates) and answer the following questions:


### 1.1 Briefly describe who you think performed this analysis. What is their expertise? What is their goal for the article? Give 3 examples of domain tasks featured in the article. (10 points)

<!-- grade_id: cell-d5cc91ff23d1fd3d, points: 0, grade: true, solution: true -->
The fact that the above article is hosted under the Chicaco Metropolitan Agency for Planning's(CMAP) governmental web domain name highly suggests that the analyst who authored the article is affiliated with the government of illinois through the aforementioned organization. The analayst can feasibly tout a degree of expertise in highway/roadway planning on account of the relative ease and accuracy with which they utilize domain specific vocabulary. Three examples of domain-specific tasks discussed in the article are:

1. What is the relationship between highway congestion and crash rates?
2. Do arterial roadways witness crashes at a different rate than highways?
3. Do highways with excess tollways/merge locations witness crashes at a different rate than highways with a more free design?

### 1.2  For each domain task describe the abstract task (10 points)

<!-- grade_id: cell-06458778598db3cb, points: 0, grade: true, solution: true -->
1. Find the correlation between variable A and variable B
2. Find the statistical difference between population A and population B
3. Find the statistical difference between population A and population B

## Part 2. Pandas programming exercise (80 points)
We have provided some code to create visualizations based on these two datasets:
1. [Historic Congestion](assets/Pulaski.small.csv.gz) 
2. [Traffic Crashes](assets/Traffic.Crashes.csv.gz)

Complete each assignment function and run each cell to generate the final visualizations


```python
import pandas as pd
import numpy as np
import altair as alt

pd.set_option('display.max_columns', None)
pd.set_option('display.max_rows', None)
```

```python
# enable correct rendering
alt.renderers.enable('default')
```

Output:
```
RendererRegistry.enable('default')
```

```python
# uses intermediate json files to speed things up
alt.data_transformers.enable('json')
```

Output:
```
DataTransformerRegistry.enable('json')
```

### PART A: Historic Congestion ( 55 points)
For parts 2.1 to 2.5 we will use the Historic Congestion dataset. This dataset contains measures of speed for different segments. For this subsample, the available measures are limited to traffic on Pulaski Road in 2018.

### 2.1 Read and resample (15 points)
Complete the `read_csv` and `get_group_first_row` functions.
Since our dataset is large we want to only grab one measurement per hour for each segment. To do this, we will resample by selecting the first measure for each month, day, hour on each segment. Complete the `get_group_first_row` function to achieve this. Note that the file we are loading is compressed--depending on how you load the file, this may or may not make a difference ([you'll want to look at the API documents](https://pandas.pydata.org/pandas-docs/stable/reference/index.html)).

<!-- grade_id: read_file, points: 0, grade: false, solution: true -->
```python
def read_csv(filename):
    """Read the csv file from filename (uncompress 'gz' if needed)
    return the dataframe resulting from reading the columns
    """
    dataframe = pd.read_csv(filename)
    return dataframe
    
    #raise NotImplementedError()
```

<!-- grade_id: read_file_test, points: 5, grade: true, solution: false -->
```python
# Save the congestion dataframe on hist_con
hist_con = read_csv('assets/Pulaski.small.csv.gz')
print(hist_con.shape)
assert hist_con.shape == (3195450, 10)
assert list(hist_con.columns) == ['TIME','SEGMENT_ID','SPEED','STREET','DIRECTION','FROM_STREET','TO_STREET',
                                  'HOUR','DAY_OF_WEEK','MONTH']
```

Output:
```
(3195450, 10)
```

<!-- grade_id: resample, points: 0, grade: false, solution: true -->
```python
def get_group_first_row(df, grouping_columns):
    """Group rows for the grouping columns and return the first row belonging to each group
    (you can look at first() for reference). We'll write this function to be more general in case
    we want to use it for a different resample.
    return a dataframe without a hierarchical index (use default index)
    
    See the example link below if you want a better sense of what this should return
    """
    dataframe = df
    dataframe_2 = dataframe.groupby(grouping_columns).first()
    dataframe_2 = dataframe_2.reset_index(level=grouping_columns)
    return dataframe_2
    #raise NotImplementedError()
```

```python
# test your code, we want segment_rows to be resampled version of hist_con where we've grouped by the
# properties month, day_of_week, hour, and segment_id and returned the first measure of each group
segment_rows = get_group_first_row(hist_con, ['MONTH','DAY_OF_WEEK', 'HOUR', 'SEGMENT_ID'])

```

The table should look something like [this](assets/segment_rows.png)

<!-- grade_id: resample_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

### 2.2 Basic Bar Chart Visualization (10 points)
We want to create a visualization for the *average speed* of each segment (across all the samples). To do this, we're going to want to group by each segment and calculate the average speed on each. Complete this code on the `average_speed_per_segment` function.

<!-- grade_id: average_speed_per_segment, points: 0, grade: false, solution: true -->
```python
def average_speed_per_segment(df):
    """Group rows by SEGMENT_ID and calculate the mean of each
    return a series where the index is the segment id and each value is the average speed per segment
    """
    dataframe = df
    dataframe = dataframe.groupby('SEGMENT_ID').mean()
    dataframe = dataframe.drop(['MONTH', 'DAY_OF_WEEK', 'HOUR'], axis=1)
    return dataframe

    #raise NotImplementedError()
```

```python
average_speed_per_segment(segment_rows)
```

Output:
```
SPEED
SEGMENT_ID           
19          12.251926
20          15.274452
21          12.141079
22          12.346769
23          12.716657
24          14.882632
25          13.076467
26          15.965027
27          21.769413
28          12.197392
29          12.384114
30          12.505631
31          11.920569
32          13.241257
33          12.955542
34          12.285714
35          13.336692
36          16.472436
37          14.086544
38          13.835803
39          15.781269
40          14.314167
41          16.032009
42          15.579135
43          17.809129
44          20.890338
45          17.608773
46          13.761707
47          14.046236
48          16.059870
49          14.053942
50          13.061055
51          13.075874
52          15.416123
53          13.803201
54          15.364552
55          14.982810
56          19.681091
57          17.820391
59          15.556609
60          18.486070
61          14.847659
62          13.339063
63          13.350919
64  
```

```python
# calculate the average speed per segment
average_speed = average_speed_per_segment(segment_rows)

# create labels for the visualization
labels = average_speed.index.astype(str)

# grab the values from the table
values = pd.DataFrame(average_speed).reset_index()

# create a chart
base = alt.Chart(values)

# we're going to "encode" the variables, more on this next assignment
encoding = base.encode(
    x= alt.X(
            'SEGMENT_ID:Q',
            title='Segment ID',
            scale=alt.Scale(zero=False)
    ),
    y=alt.Y(
            'sum(SPEED):Q',
            title='Speed Average MPH'
    ),
)

# we're going to use a bar chart and set various parameters (like bar size and title) to make it readable
encoding.mark_bar(size=7).properties(title='Average Speed per Segment',height=300, width=900)
```

Output:
```
alt.Chart(...)
```

The table should look something like [this](assets/average_speed.png)

<!-- grade_id: average_speed_per_segment_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

### 2.3 Create a basic pivot table (10 points)
For the next visualization, we need a more complex transformation that will allow us to see the average speed for each month. To do this, we will create a pivot table where the index is the month, and each column is a segment id. We will put the average speed in the cells. From the table, we'll be able to find the month (by index)--giving us the row, and pick the column corresponding to the segment we care about.

Complete the `create_pivot_table` function for this

<!-- grade_id: create_pivot_table, points: 0, grade: false, solution: true -->
```python
def create_pivot_table (df):
    """return a pivot table where:
    each row i is a month
    each column j is a segment id
    each cell value is the average speed for the month i in the segment j
    """
    dataframe = df
    table = pd.pivot_table(dataframe, values='SPEED', index=['MONTH'],
                    columns=['SEGMENT_ID'], aggfunc=np.mean)
    return table
    
    #raise NotImplementedError()
```

```python
# run the code and see what's in the table
pivot_table = create_pivot_table(segment_rows)
pivot_table
```

Output:
```
SEGMENT_ID         19         20         21         22         23         24  \
MONTH                                                                          
2            6.857143  16.142857  13.571429  19.571429  18.285714  15.857143   
3           10.773810  14.863095  11.696429  11.815476  13.583333  16.244048   
4           11.744048  14.958333  11.791667  12.071429  13.208333  16.779762   
5           11.357143  14.738095  11.369048  11.916667  12.023810  13.220238   
6           11.630952  14.583333  13.011905  12.279762  12.428571  14.678571   
7           11.755952  13.595238  10.880952  12.238095  12.267857  14.321429   
8           12.988095  15.446429  12.303571  13.315476  13.023810  15.827381   
9           13.970238  17.059524  14.398810  13.452381  12.017857  14.869048   
10          13.708333  15.666667  12.434524  13.041667  12.422619  13.714286   
11          12.970238  16.107143  11.922619  11.476190  12.125000  15.607143   
12          11.845238  15.690476  11.541
```

The table should look something like [this](assets/pivot_table.png)

```python
# we're going to implement a transformation to put the pivot table into a 'long form' because it
# is easier to specify the visualization. You can print out hm_pivot_table to see what it looks like
hm_pivot_table = pivot_table.copy().unstack().reset_index()
hm_pivot_table['SPEED'] = hm_pivot_table[0]
hm_pivot_table.drop(0,axis=1,inplace=True)

# create the visualization. We're going to use rectangles (a heat map of sorts). We'll use the segment_id to
# figure out the horizontal placement (x), the month as the vertical (y) and use color to encode the speed.
encoding = alt.Chart(hm_pivot_table).mark_rect().encode(
    x='SEGMENT_ID:O',
    y='MONTH:O',
    color='SPEED:Q'
)

encoding.properties(title='Average Speed per Segment per Month',height=300, width=800)
```

Output:
```
alt.Chart(...)
```

<!-- grade_id: create_pivot_table_test, points: 10, grade: true, solution: false -->
```python
# test function
pivot_table = create_pivot_table(segment_rows)
# check that the rows are months and columns are segments
assert pivot_table.shape == (11, 78), "Problem 2.3, first test"
# check that the value is the average
assert int(pivot_table.loc[2,19]) == 6, "Problem 2.3, second test"
assert int(pivot_table.loc[3,19]) == 10, "Problem 2.3, third test"
```

### 2.4 Sorting, Transforming, and Filtering (20 points)
Without telling you too much about the visualization we want to create next (that's part of the bonus below), we need to get the data into a form we can use. 
- We're going to need to sort the dataframe by one or more columns (this is the `sort` function). 
- We'll want to create a derivative column that is the time of the measurement rounded to the nearest hour (`time_to_hours`)
- We need to "facet" the data into groups to generate different visualizations. 
- We need a function that selects part of the dataframe that matches a specific characteristic (`filter_orientation`)
- Grab a specific column from the dataframe (`select_column`)

<!-- grade_id: sort, points: 0, grade: false, solution: true -->
```python
def sort(df, sorting_columns):
    """Sort the rows by the columns
    return the sorted dataframe
    """
    dataframe = df
    dataframe = dataframe.sort_values(sorting_columns)
    return dataframe
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
segment_rows = sort(segment_rows, ['SEGMENT_ID'])
```

<!-- grade_id: sort_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<!-- grade_id: time_to_hours, points: 0, grade: false, solution: true -->
```python
def time_to_hours(df):
    """ Add a column (called TIME_HOURS) based on the data in the TIME column and rounded up 
    the value to the nearest hour.  For example, if the original TIME row said: 
    02/28/2018 05:40:00 PM we want 2018-02-28 18:00:00  
    (the change is that 5:40pm was rounded up to 6:00pm and the TIME_HOUR column is 
    actually a proper datetime and not a string).
    """
    dataframe = df
    dataframe['TIME_HOURS'] = pd.to_datetime(dataframe['TIME']).dt.round('H')
    #dataframe['TIME_HOURS'] = dataframe['TIME_HOURS'].dt.round('H')
    return dataframe
    #raise NotImplementedError()
```

```python
segment_rows = time_to_hours(segment_rows)

```

<!-- grade_id: time_to_hours_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<!-- grade_id: filter_orientation, points: 0, grade: false, solution: true -->
```python
def filter_orientation(df, traffic_orientation):
    """ Filter the rows according to the traffic orientation
    return a df that is a subset of the original with the desired orientation
    """
    dataframe = df
    dataframe_subset = dataframe[dataframe['DIRECTION'] == traffic_orientation]
    return dataframe_subset
    #raise NotImplementedError()
```

```python
sb = filter_orientation(segment_rows, 'SB')
nb = filter_orientation(segment_rows, 'NB')

```

The sb table should look something like [this](assets/sb.png)

<!-- grade_id: filter_orientation_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<!-- grade_id: select_column, points: 0, grade: false, solution: true -->
```python
def select_column(df, column_name):
    """ Select a column from the df
    return a series with the desired column
    """
    dataframe = df
    dataframe_selected_column = dataframe[column_name]
    return dataframe_selected_column
    #raise NotImplementedError()
```

<!-- grade_id: select_column_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

```python
# we're going to remove speeds of -1 (no data)
sb = sb[sb.SPEED > -1]
nb = nb[nb.SPEED > -1]
```

```python

alt.data_transformers.disable_max_rows()
alt.Chart(sb.append(nb)).mark_rect().encode(
    x='month(TIME_HOURS):T',
    y='FROM_STREET:N',
    color='mean(SPEED):Q',
    facet='DIRECTION:N'
).properties(
    width=300,
    height=400
)

```

Output:
```
alt.Chart(...)
```

### 2.5 (Bonus)  Traffic heatmap visualization (up to 2 points)
Looking at the visualization above (the one showing Northbound versus Southbound facets), what domain/abstract tasks are fulfilled by this visualization? List at least one domain task and the corresponding abstract task.

<!-- grade_id: complex_heatmap_vis, points: 0, grade: true, solution: true -->
Domain Task: Is there a relationship between time and traffic congestion on northbound lanes and southbound lanes, and is there a significant difference between both populations?

Abstract Tasks: (1) Find the correlation between variable A and variable B
                (2) Find the statistical difference between population A and population B

### PART B: Crashes (25 points)
For parts 2.6 and 2.7 we will use the Crashes dataset. This dataset contains crash entries recording the time of the accident, the street, and the street number where the accident occurred. You will work with accidents recorded on Pulaski Road


```python
crashes = read_csv('assets/Traffic.Crashes.csv.gz')
crashes_pulaski = crashes[crashes.STREET_NAME == 'PULASKI RD']
```

```python
crashes_pulaski['STREET_NO'].max()
```

Output:
```
11480
```

### 2.6 Calculate summary statistics for grouped streets (15 points)
- Group the streets every 300 units (street numbers). Hint: You can use the `pd.cut` function
- Calculate the number of accidents (count rows) and the total of injuries (sum injuries total) for each of these 300-chunk road segments. Do this *for each direction*.

Complete bin_crashes and calculate_group_aggregates functions for this

<!-- grade_id: bin_crashes, points: 0, grade: false, solution: true -->
```python
def bin_crashes(df):
    """ Assign each crash instance a category (bin) every 300 house number units starting from 0
    Return a new dataframe with a column called BIN where each value is the start of the bin
    i.e. 0 is the label for records with street number n, such that 1 <= n <= 300
    300 is the label for records with with n at 301 <= n <= 600, and so on.
    """
    bin_values = np.arange(0,12000,300)
    df['BIN'] = pd.cut(df['STREET_NO'],bin_values,labels=bin_values[0:-1])
    return df
    #raise NotImplementedError()
```

```python
binned_df = bin_crashes(crashes_pulaski)
```

```python
binned_df.head()
```

Output:
```
RD_NO CRASH_DATE_EST_I              CRASH_DATE  POSTED_SPEED_LIMIT  \
6    JC100005              NaN  12/31/2018 11:45:00 PM                  25   
51   JB574321              NaN  12/31/2018 09:20:00 PM                  30   
126  JB574112              NaN  12/31/2018 05:30:00 PM                  30   
133  JC111739              NaN  12/31/2018 05:30:00 PM                  30   
240  JB573809              NaN  12/31/2018 01:30:00 PM                  30   

    TRAFFIC_CONTROL_DEVICE      DEVICE_CONDITION WEATHER_CONDITION  \
6           TRAFFIC SIGNAL  FUNCTIONING PROPERLY           UNKNOWN   
51          TRAFFIC SIGNAL  FUNCTIONING PROPERLY              RAIN   
126            NO CONTROLS           NO CONTROLS              RAIN   
133            NO CONTROLS           NO CONTROLS              RAIN   
240            NO CONTROLS           NO CONTROLS              RAIN   

         LIGHTING_CONDITION      FIRST_CRASH_TYPE  \
6    DARKNESS, LIGHTED ROAD               TURNING   
51   DARKNES
```

A sample of the relevant columns from the table would look something like [this](assets/binned_df.png). We can also create a histogram of street numbers to see which are the most prevalent. It should look something like [this](assets/street_no.png).

```python
# create this vis
alt.Chart(binned_df).mark_bar().encode(
    alt.X('BIN'),
    alt.Y('count()')
)
```

Output:
```
alt.Chart(...)
```

<!-- grade_id: bin_crashes_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<!-- grade_id: calculate_group_aggregates, points: 0, grade: false, solution: true -->
```python
def calculate_group_aggregates(df):
    """ 
    Return a df with the count of accidents in a 'ACCIDENT_COUNT' column and 'INJURIES_SUM'
    """
    dataframe = df
    grouping_columns = ['BIN', 'STREET_DIRECTION']
    dataframe = dataframe.groupby(grouping_columns)['INJURIES_TOTAL'].agg({'ACCIDENT_COUNT': 'count', 'INJURIES_SUM': 'sum'})
    dataframe = dataframe.reset_index(level=grouping_columns)
    
    return dataframe
    #raise NotImplementedError()
```

```python
aggregates = calculate_group_aggregates(binned_df)
aggregates.head(5)

```

Output:
```
BIN STREET_DIRECTION  ACCIDENT_COUNT  INJURIES_SUM
0    0                N            52.0          18.0
1    0                S            44.0          17.0
2  300                N            37.0          19.0
3  300                S            49.0          17.0
4  600                N            69.0          19.0
```

The table should look something like [this](assets/aggregates.png)

<!-- grade_id: calculate_group_aggregates_TEST, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

Just for fun, here's a plot of injuries in the North and South directions based on bin. This may also help you debug your code. Depending on whether you removed N/A or if you hardcoded things, you may see slight differences. Here's what it might  [look like](assets/direction_injuries.png)

```python
alt.Chart(aggregates).mark_point().encode(
    alt.Color('STREET_DIRECTION'),
    alt.X('BIN'),
    alt.Y('INJURIES_SUM')
)
```

Output:
```
alt.Chart(...)
```

### 2.7 Sort the street ranges (10 points)
- Sort the dataframe so North streets are in descending order and South streets are in ascending order
- You are provided with a 'sort' arrray that contains this desired order. Use a categorical (pd.Categorial) column to order the dataframe according to this array.

<!-- grade_id: categorical_sorting, points: 0, grade: false, solution: true -->
```python
crashed_range = list(range(0, crashes_pulaski.STREET_NO.max()+1000, 300))
sort = ['N ' + str(s) for s in crashed_range[::-1]] + ['S ' + str(s) for s in crashed_range]
def categorical_sorting(df, sort):
    """ Create a column called ORDER_LABEL that contains a concatenation of the street direction and the street range
    Set the sort order of this column to the provided sort array (the elements of this column should be in the same order
    of the array)
    Sort the dataframe by this column
    """
    dataframe = df
    dataframe['BIN'] = dataframe['BIN'].astype(str)
    dataframe['STREET_DIRECTION'] = dataframe['STREET_DIRECTION'].astype(str)
    dataframe['ORDER_LABEL'] = dataframe['STREET_DIRECTION'] + ' ' +  dataframe['BIN']
    dataframe['ORDER_LABEL'] = pd.Categorical(dataframe['ORDER_LABEL'], 
                        categories=sort,
                        ordered=True)
                                                        
    dataframe.sort_values('ORDER_LABEL', inplace=True)                                                 

    return dataframe
    #raise NotImplementedError()
```

```python
sorted_groups = categorical_sorting(aggregates, sort)

```

The table should look something like [this](assets/sorted_groups.png)

<!-- grade_id: categorical_sorting_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

Again, just for kicks, let's see where injuries happen. We're going to color bars by the bin and preserve our ascending/descending visualization. We can probably imagine other (better) ways to visualize this data, but this may be useful for you to debug. The visualization should look something like [this](assets/order_injuries.png)

```python
alt.Chart(sorted_groups).mark_bar().encode(
    alt.X('ORDER_LABEL:O', sort=sort),
    alt.Y('INJURIES_SUM:Q'),
    alt.Color('BIN:Q')
).properties(
    width=400
)
```

Output:
```
alt.Chart(...)
```

Ok, let's actually make a useful visualization using some of the dataframes we've created. As a bonus, we're going to ask you what you would use this for.

```python
# to make the kind of chart we are interested in we're going to build it out of three different charts and
# put them together at the end

# this is going to be the left chart
bar_sorted_groups = sorted_groups[['ACCIDENT_COUNT','INJURIES_SUM']].unstack().reset_index() \
    .rename({'level_0':'TYPE','level_1':'SPEED',0:'COUNT'},axis=1)

a = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'ACCIDENT_COUNT').encode(
    x=alt.X('COUNT:Q',sort='descending'),
    y=alt.Y('SPEED:O',axis=None),
    color=alt.Color('TYPE:N', 
                    legend=None,
                    scale=alt.Scale(domain=['ACCIDENT_COUNT', 'INJURIES_SUM'],
                                    range=['blue', 'orange']))
).properties(
    title='ACCIDENT_COUNT',
    width=300,
    height=600
)

# middle "chart" which actually won't be a chart, just a bunch of labels
b = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'ACCIDENT_COUNT').encode(
    y=alt.Y('SPEED:O', axis=None),
    text=alt.Text('SPEED:Q')
).mark_text().properties(title='SPEED',
                         width=20,
                         height=600)

# and the right most chart
c = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'INJURIES_SUM').encode(
    x='COUNT:Q',
    y=alt.Y('SPEED:O',axis=None),
    color=alt.Color('TYPE:N', 
                    legend=None,
                    scale=alt.Scale(domain=['ACCIDENT_COUNT', 'INJURIES_SUM'],
                                    range=['blue', 'orange']))
).properties(
    title='INJURIES_SUM',
    width=300,
    height=600
)

# put them all together 

a | b | c

```

Output:
```
alt.HConcatChart(...)
```

## 2.8 (Bonus) Accident barchart visualization (up to 2 points)
Looking at the visualization we generated above (part 2.7), what domain/abstract tasks are fulfilled by this visualization? List at least one domain task and the corresponding abstract task. 

<!-- grade_id: complex_barchart_vis, points: 0, grade: true, solution: true -->
Domain Task: What is the relationship between speed and accident count, and speed and total injuries?

Abstract Task: Find the correlation between variable A and variable B